## HW02: Exploratory Data Analysis (Part 01)

This is Assignment 02 for the course "Introduction to Data Science" at the Faculty of Information Technology, University of Science, Vietnam National University, Ho Chi Minh City.

---

(Latest update: 26/10/2024)

Student Name: Võ Hùng Khoa

Student ID: 22127202

---

## **How to Complete and Submit the Assignment**

&#9889; **Note**: You should follow the instructions below. If anything is unclear, you need to contact the teaching assistant or instructor immediately for timely support.

**How to Do the Assignment**

You will work directly on this notebook file. First, fill in your full name and student ID (MSSV) in the header section of the file above. In the file, complete the tasks in sections marked:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
Or for optional code sections:
```python
# YOUR CODE HERE (OPTION)
```
For markdown cells, complete the answer in the section marked:
```markdown
YOUR ANSWER HERE
```

**How to Submit the Assignment**

Before submitting, select `Kernel` -> `Restart Kernel & Run All Cells` if you are using a local environment, or `Runtime -> Restart session` and run all if using Google Colab, to ensure everything works as expected.

Next, create a submission folder with the following structure:
- Folder named `MSSV` (for example, if your student ID is `1234567`, name the folder `1234567`)
    - File `HW02.ipynb` (no need to submit other files)

Finally, compress this `MSSV` folder in `.zip` format (not `.rar` or any other format) and submit it via the link on Moodle.\
<font color=red>Please make sure to strictly follow this submission guideline.</font>

---

## Import

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
%matplotlib inline
```

---

## Data collecting

The two datasets are related to red and white variants of the Portuguese "Vinho Verde" wine. For more details, consult: http://www.vinhoverde.pt/en/ or the reference [Cortez et al., 2009].  Due to privacy and logistic issues, only physicochemical (inputs) and sensory (the output) variables are available (e.g. there is no data about grape types, wine brand, wine selling price, etc.). You can download these data from this [link](https://archive.ics.uci.edu/dataset/186/wine+quality).

These datasets can be viewed as classification or regression tasks.  The classes are ordered and not balanced (e.g. there are many more normal wines than excellent or poor ones). Outlier detection algorithms could be used to detect the few excellent or poor wines. Also, we are not sure if all input variables are relevant. So it could be interesting to test feature selection methods.

---

## Exploring your data

### Read raw data from file (0.25 points)

<!-- grade_id: cell-a1f731c2bc3e6167, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
raw_white = pd.read_csv("data/winequality-white.csv", sep=";")
raw_red = pd.read_csv("data/winequality-red.csv", sep=";")

```

<!-- grade_id: cell-f164df2170af969e, points: 0.15, grade: true, solution: false -->
```python
# TEST 1
raw_white.head()
```

Output:
```
fixed acidity  volatile acidity  citric acid  residual sugar  chlorides  \
0            7.0              0.27         0.36            20.7      0.045   
1            6.3              0.30         0.34             1.6      0.049   
2            8.1              0.28         0.40             6.9      0.050   
3            7.2              0.23         0.32             8.5      0.058   
4            7.2              0.23         0.32             8.5      0.058   

   free sulfur dioxide  total sulfur dioxide  density    pH  sulphates  \
0                 45.0                 170.0   1.0010  3.00       0.45   
1                 14.0                 132.0   0.9940  3.30       0.49   
2                 30.0                  97.0   0.9951  3.26       0.44   
3                 47.0                 186.0   0.9956  3.19       0.40   
4                 47.0                 186.0   0.9956  3.19       0.40   

   alcohol  quality  
0      8.8        6  
1      9.5        6  
2     10.1        6  
3
```

<!-- grade_id: cell-37ca4dbd9cc88d0a, points: 0.1, grade: true, solution: false -->
```python
# TEST 2
raw_red.head()
```

Output:
```
fixed acidity  volatile acidity  citric acid  residual sugar  chlorides  \
0            7.4              0.70         0.00             1.9      0.076   
1            7.8              0.88         0.00             2.6      0.098   
2            7.8              0.76         0.04             2.3      0.092   
3           11.2              0.28         0.56             1.9      0.075   
4            7.4              0.70         0.00             1.9      0.076   

   free sulfur dioxide  total sulfur dioxide  density    pH  sulphates  \
0                 11.0                  34.0   0.9978  3.51       0.56   
1                 25.0                  67.0   0.9968  3.20       0.68   
2                 15.0                  54.0   0.9970  3.26       0.65   
3                 17.0                  60.0   0.9980  3.16       0.58   
4                 11.0                  34.0   0.9978  3.51       0.56   

   alcohol  quality  
0      9.4        5  
1      9.8        5  
2      9.8        5  
3
```

Combine the two dataframes and add a `colors` column to distinguish between red and white wines. Named the dataframe as `raw_df`.

```python
raw_white["colors"] = "white"
raw_red["colors"] = "red"
raw_df = pd.concat([raw_white, raw_red], ignore_index=True)
```

### How many rows and how many columns does the raw data have? (0.25 points)

Next, you calculate the number of rows and columns of two DataFrames `raw_white` and `raw_red`. Then, store it in the variable `shape_white` and `shape_red`(tuple). 

<!-- grade_id: cell-e4e95aeda5e3d712, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
shape_white = raw_white.shape
shape_red = raw_red.shape
shape = raw_df.shape
```

<!-- grade_id: cell-11c8f187fa324c33, points: 0.25, grade: true, solution: false -->
```python
# TEST
assert shape_white == (4898, 13)
assert shape_red == (1599, 13)
assert shape == (6497, 13)
```

### What does each line mean? Does it matter if the lines have different meanings?

**What does each line mean?**  
Each line represents a single wine sample, containing its physicochemical properties (input variables) and the sensory quality score (output variable).  


**Does it matter if the lines have different meanings?**  
Yes, consistency in the meaning of lines is crucial. If the lines represent different contexts (e.g., varying wine types or evaluation methods), the model's predictions may become inaccurate due to mixed data patterns.

### Does the raw data have duplicate rows? (0.25 points)

Next, you calculate the number of rows with duplicate indexes and store it in the variable `red_num_duplicated_rows` and `white_num_duplicated_rows`. In a group of lines with the same index, the first line is not counted as repeated.

<!-- grade_id: cell-f3dec99d2782b78c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
# calculate the number of rows with duplicate indexes and store it in the variable `red_num_duplicated_rows` and `white_num_duplicated_rows In a group of lines with the same index, the first line is not counted as repeated. 

# Count duplicate rows by index for red wine data
red_num_duplicated_rows = raw_red.duplicated(keep='first').sum()

# Count duplicate rows by index for white wine data
white_num_duplicated_rows = raw_white.duplicated(keep='first').sum()

```

<!-- grade_id: cell-14c94f9e78ffc028, points: 0.25, grade: true, solution: false -->
```python
# TEST
assert white_num_duplicated_rows == 937
assert red_num_duplicated_rows == 240
```

```python
raw_red.drop_duplicates(keep=False, inplace=True)
raw_white.drop_duplicates(keep=False, inplace=True)
raw_df.drop_duplicates(keep=False, inplace=True)
```

### What does each column mean? (0.25đ)

Read file `winequality.names` and indentify what's type (numerical/ categorical; continous/ discrete; ordinal/ nominal) for each attributes in the dataset.

<!-- grade_id: cell-391565b23a9d6966, points: 0.25, grade: true, solution: true -->
<center>

| **Variable**            | **Type**                   |
|--------------------------|----------------------------|
| Fixed acidity            | Numerical, Continuous     |
| Volatile acidity         | Numerical, Continuous     |
| Citric acid              | Numerical, Continuous     |
| Residual sugar           | Numerical, Continuous     |
| Chlorides                | Numerical, Continuous     |
| Free sulfur dioxide      | Numerical, Discrete       |
| Total sulfur dioxide     | Numerical, Discrete       |
| Density                  | Numerical, Continuous     |
| pH                       | Numerical, Continuous     |
| Sulphates                | Numerical, Continuous     |
| Alcohol                  | Numerical, Continuous     |
| Quality                  | Categorical, Ordinal      |
| Color                    | Categorical, Nominal      |

<center>

### What data type does each column currently have? Are there any columns whose data types are not suitable for further processing? (0.25 points)

Next, you calculate the data type (dtype) of each column in DataFrame `raw_white` and save the result into Series `white_dtypes` (This Series has the index as the column name). Do similar for `raw_red`.

<!-- grade_id: cell-4842c5873ccd8cab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
white_dtypes = raw_white.dtypes
red_dtypes = raw_red.dtypes
```

<!-- grade_id: cell-ce4b8d36633948ce, points: 0.15, grade: true, solution: false -->
```python
# TEST 1
white_dtypes
```

Output:
```
fixed acidity           float64
volatile acidity        float64
citric acid             float64
residual sugar          float64
chlorides               float64
free sulfur dioxide     float64
total sulfur dioxide    float64
density                 float64
pH                      float64
sulphates               float64
alcohol                 float64
quality                   int64
colors                   object
dtype: object
```

<!-- grade_id: cell-e89ce01522920567, points: 0.1, grade: true, solution: false -->
```python
# TEST 2
red_dtypes

# We can see that quality column is not suitable type, it should be in categorical type 
```

Output:
```
fixed acidity           float64
volatile acidity        float64
citric acid             float64
residual sugar          float64
chlorides               float64
free sulfur dioxide     float64
total sulfur dioxide    float64
density                 float64
pH                      float64
sulphates               float64
alcohol                 float64
quality                   int64
colors                   object
dtype: object
```

### For each column with numeric data type, how are the values distributed? (1đ)

For columns with numeric data types, you will calculate:
- Percentage (from 0 to 100) of missing values
- The min
- The lower quartile (phân vị 25)
- The median (phân vị 50)
- The upper quartile (phân vị 75)
- The max

You will save the results to a DataFrame `num_col_info_red` and `num_col_info_white`, where:
- The names of the columns are the names of the numeric columns in `raw_red` and `raw_white`
- Names of rows: "missing_ratio", "min", "lower_quartile", "median", "upper_quartile", "max"  

For ease of viewing, you round all values to 1 decimal place using the `.round(1)` method.

<!-- grade_id: cell-bac7c3dae96f89ba, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
def calculate_numeric_stats(df):

    stats = {
        # Calculate the missing ratio of the numeric columns
        "missing_ratio": df.isnull().mean() * 100,

        # Calculate the min of the numeric columns
        "min": df.min(),

        # Calculate the lower_quartile of the numeric columns
        "lower_quartile": df.quantile(0.25),

        # Calculate the median of the numeric columns
        "median": df.median(),

        # Calculate the upper_quartile of the numberic columns
        "upper_quartitle": df.quantile(0.75),

        # Calculate the max of the numeric columns
        "max": df.max()
    }
    return pd.DataFrame(stats).round(1).transpose()
    
# Calculate the numeric statistics for the given DataFrame
num_col_info_red = calculate_numeric_stats(raw_red.select_dtypes(include=[np.number]))
num_col_info_white = calculate_numeric_stats(raw_white.select_dtypes(include=[np.number]))

# Rename the columns to make them more descriptive
num_col_info_red.columns = raw_red.select_dtypes(include=[np.number]).columns
num_col_info_white.columns = raw_white.select_dtypes(include=[np.number]).columns


# Define some statistical functions
def missing_ratio(series):
    return series.isnull().mean()*100

def lower_quartile(series):
    return series.quantile(0.25)

def median(series):
    return series.median()

def upper_quartile(series):
    return series.quantile(0.75)


```

<!-- grade_id: cell-8f6bde86dbc648e0, points: 1, grade: true, solution: false -->
```python
# TEST
num_col_info_red = num_col_info_red.agg([missing_ratio, "min", lower_quartile, median, upper_quartile, "max"])
num_col_info_white = num_col_info_white.agg([missing_ratio, "min", lower_quartile, median, upper_quartile, "max"])
```

### For each column with a non-numeric data type, how are the values distributed?(1đ)

For columns with non-numeric data types, you calculate:
- Percentage (from 0 to 100) of missing values
- Number of values (the values here are different values and we do not consider missing values): with columns whose type is categorical, it is a set with a finite number of categories. Directly counting the number of values in these columns doesn't make much sense, so it's better to count the number of elements of all types.
- The percentage (from 0 to 100) of each value is sorted by decreasing percentage (we do not consider missing values, the ratio is the ratio compared to the number of non-missing values): you use a dictionary to store , key is the value, value is the percentage; With the column corresponding to each type, the method is similar to above.

You will save the results to DataFrame `cat_col_info_df`, where:
- The names of the columns are the names of the non-numeric columns in `raw_df`
- The names of the lines are: "missing_ratio", "num_values", "value_ratios"

For ease of viewing, you round all values to 1 decimal place using the `.round(1)` method.

```python
len(raw_df.select_dtypes(exclude=np.number).dropna())
```

Output:
```
2
```

<!-- grade_id: cell-26c7e877fbb463bf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
def cat_col_info_df(df):
    stats = {
        # Calculate the missing ratio of the numeric columns
        "missing_ratio": df.isnull().mean() * 100,

        # Calculate the min of the numeric columns
        "num_values": len(df.value_counts()),

        


        "value_ratios": 
    }
    return pd.DataFrame(stats).round(1).transpose()
```

<!-- grade_id: cell-e25b38776bb99db3, points: 1, grade: true, solution: false -->
```python
# TEST
cat_col_info_df = cat_col_info_df.agg([missing_ratio, num_values, value_ratios])
cat_col_info_df
```

Output:
```
AttributeError                            Traceback (most recent call last)
Cell In[18], line 2
      1 # TEST
----> 2 cat_col_info_df = cat_col_info_df.agg([missing_ratio, num_values, value_ratios])
      3 cat_col_info_df

AttributeError: 'function' object has no attribute 'agg'
```

## Exploring your data (Cont)

### Task 01: Analyzing wine quality via histogram.

<!-- grade_id: cell-e45473244e9e6f86, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Task 02: Aanalyzing acidity in red and white wine.

<!-- grade_id: cell-b052e7de9514401d, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e1412c0b663cae47, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-189d0d108d8a7505, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Task 03: Analyzing sulfur dioxide for each type of wines. Try to focus on the ratio of free and total sulfur.

<!-- grade_id: cell-ccdb9ef80f167bdf, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3de6ed5efea228c7, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2e8c6bd101daefc2, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Task 04: What about the residual sugar? How does it affect wine quality?

<!-- grade_id: cell-6141cade9bdb3832, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Task 05: Alcohol plays a vital role in wine production. How does it affect wine quality?

<!-- grade_id: cell-ca8827b632fdd2df, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Task 06: For each type of wine, how is density distributed? Is it importance to wine quality?

<!-- grade_id: cell-a5afb0e5f433a1f8, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Task 07: How are Chloride characteristics distributed in these two types of wines? 

<!-- grade_id: cell-cf2a2cac703189b2, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Analyze the correlation between two or more variables

### Task 08: Correlation matrix analysis

<!-- grade_id: cell-bfac6ff937ff1ab9, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Analysis of the influence of variables on wine quality

#### Task 09: Investigating the relationship between `Density` and `Quality`

<!-- grade_id: cell-f2bf686379e5256c, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-435d26d45f12207b, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 10: Investigating the relationship between `Alcohol` and `Quality`

<!-- grade_id: cell-9f4c1e9505987aae, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 11: Investigating the relationship between `Chlorides` and `Quality`

<!-- grade_id: cell-7eb72396937094db, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 12: Investigating the relationship between `Volatile Acidity` and `Quality`

<!-- grade_id: cell-83e1238956bda93c, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 13: Investigation of the relationship between total SO2 and residual sugar after fermentation

<!-- grade_id: cell-95dc80e4d9cc7ac2, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 14: Investigation of the impact of alcohol concentration and sugar content on alcohol density

<!-- grade_id: cell-18f74d0dae43e399, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e935af33e66b4d63, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Analysis based on wine color

#### Task 15: Analyze the relationship between color and density of wine

<!-- grade_id: cell-9bd179f1ef0bc6a8, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 16: Analyze the relationship between color and residual sugar after fermentation

<!-- grade_id: cell-822d5c85c49a67bc, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 17: Analyze the relationship between color and total sulfur content in wine

<!-- grade_id: cell-9e23e178877bcb14, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 18: Analyze the relationship between color and free sulfur content in wine

<!-- grade_id: cell-d45295de28543279, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 19: Analyze the relationship between color and acidity of wine

<!-- grade_id: cell-0ca188ff25f80070, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Correlation analysis between variables based on color

#### Task 20: Investigate the correlation between density and wine quality based on color

<!-- grade_id: cell-758d7c4f1bee360c, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 21: Investigate the correlation between alcohol concentration and wine quality based on color

<!-- grade_id: cell-e1cbb0e7a2366906, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 22: Investigate the correlation between salt content and wine quality based on color

<!-- grade_id: cell-fc043b0f56a7b3a6, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 23: Investigate the correlation between acidity and wine quality based on color

<!-- grade_id: cell-55e8b921bf99b6c7, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Multi-attributes analysis

#### Task 24: Chlorides and Sulphates on wine quality

<!-- grade_id: cell-7cf96ac1224bb022, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-48afbbbe3a49d417, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 25: Sulfur Dioxide Ratio and Volatile Acidity Ratio on wine quality

<!-- grade_id: cell-1879ef370d511826, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7501f4d43cff2294, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Task 26: Volatile Acidity and Alcohol on wine quality

<!-- grade_id: cell-69e64b640a50e69a, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Main Factors Influencing Wine Quality

Based on the above analysis, you should decide what factors affect wine quality.

<!-- grade_id: cell-1906324a227a5e65, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```